Prophage-encoded Hm-oscar gene recapitulates Wolbachia-induced male killing in the tea tortrix moth Homona magnanima

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public review):

Summary:

Insects and their relatives are commonly infected with microbes that are transmitted from mothers to their offspring. A number of these microbes have independently evolved the ability to kill the sons of infected females very early in their development; this male killing strategy has evolved because males are transmission dead-ends for the microbe. A major question in the field has been to identify the genes that cause male killing and to understand how they work. This has been especially challenging because most male-killing microbes cannot be genetically manipulated. This study focuses on a male-killing bacterium called Wolbachia. Different Wolbachia strains kill male embryos in beetles, flies, moths, and other arthropods. This is remarkable because how sex is determined differs widely in these hosts. Two Wolbachia genes have been previously implicated in male-killing by Wolbachia: oscar (in moth male-killing) and wmk (in fly male-killing). The genomes of some male-killing Wolbachia contain both of these genes, so it is a challenge to disentangle the two.

This paper provides strong evidence that oscar is responsible for male-killing in moths. Here, the authors study a strain of Wolbachia that kills males in a pest of tea, Homona magnanima. Overexpressing oscar, but not wmk, kills male moth embryos. This is because oscar interferes with masculinizer, the master gene that controls sex determination in moths and butterflies. Interfering with the masculinizer gene in this way leads the (male) embryo down a path of female development, which causes problems in regulating the expression of genes that are found on the sex chromosomes.

We would like to thank you for evaluating our manuscript.

Strengths:

The authors use a broad number of approaches to implicate oscar, and to dissect its mechanism of male lethality. These approaches include: a) overexpressing oscar (and wmk) by injecting RNA into moth eggs, b) determining the sex of embryos by staining female sex chromosomes, c) determining the consequences of oscar expression by assaying sex-specific splice variants of doublesex, a key sex determination gene, and by quantifying gene expression and dosage of sex chromosomes, using RNASeq, and d) expressing oscar along with masculinizer from various moth and butterfly species, in a silkmoth cell line. This extends recently published studies implicating oscar in male-killing by Wolbachia in Ostrinia corn borer moths, although the Homona and Ostrinia oscar proteins are quite divergent. Combined with other studies, there is now broad support for oscar as the male-killing gene in moths and butterflies (i.e. order Lepidoptera). So an outstanding question is to understand the role of wmk. Is it the master male-killing gene in insects other than Lepidoptera and if so, how does it operate?

We would like to thank you for evaluating our manuscript. Our data demonstrated that Oscar homologs play important roles in male-killing phenotypes in moths and butterflies; however, the functional relevance of wmk remains uncertain. As you noted, whether wmk acts as a male-killing gene in insects such as flies and beetles—or even in certain lepidopteran species—requires further investigation using diverse insect models, which we are eager to explore in future research.

Weaknesses:

I found the transfection assays of oscar and masculinizer in the silkworm cell line (Figure 4) to be difficult to follow. There are also places in the text where more explanation would be helpful for non-experts.

Thank you for your suggestion. We have revised the section on the cell-based experiment. Further, we revised the manuscript to make it accessible to a broader audience. We believe these revisions have significantly improved the clarity and comprehensiveness of our manuscript.

Reviewer #2 (Public review):

Summary:

Wolbachia are maternally transmitted bacteria that can manipulate host reproduction in various ways. Some Wolbachia induce male killing (MK), where the sons of infected mothers are killed during development. Several MK-associated genes have been identified in Homona magnanima, including Hm-oscar and wmk-1-4, but the mechanistic links between these Wolbachia genes and MK in the native host are still unclear.

In this manuscript, Arai et al. show that Hm-oscar is the gene responsible for Wolbachia-induced MK in Homona magnanima. They provide evidence that Hm-Oscar functions through interactions with the sex determination system. They also found that Hm-Oscar disrupts sex determination in male embryos by inducing female-type dsx splicing and impairing dosage compensation. Additionally, Hm-Oscar suppresses the function of Masc. The manuscript is well-written and presents intriguing findings. The results support their conclusions regarding the diversity and commonality of MK mechanisms, contributing to our understanding of the mechanisms and evolutionary aspects of Wolbachia-induced MK.

We would like to thank you for evaluating our manuscript.

Comments on revisions:

The authors have already addressed the reviewer's concerns.

We would like to thank you for evaluating our manuscript.

Reviewer #3 (Public review):

Summary:

Overall, this is a clearly written manuscript with nice hypothesis testing in a non-model organism that addresses the mechanism of Wolbachia-mediated male killing. The authors aim to determine how five previously identified male-killing genes (encoded in the prophage region of the wHm Wolbachia strain) impact the native host, Homona magnanima moths. This work builds on the authors' previous studies in which

(1) they tested the impact of these same wHm genes via heterologous expression in Drosophila melanogaster

(2) also examined the activity of other male-killing genes (e.g., from the wFur Wolbachia strain in its native host: Ostrinia furnacalis moths).

Advances here include identifying which wHm gene most strongly recapitulates the male-killing phenotype in the native host (rather than in Drosophila), and the finding that the Hm-Oscar protein has the potential for male-killing in a diverse set of lepidopterans, as inferred by the cell-culture assays.

We would like to thank you for evaluating our manuscript.

Strengths:

Strengths of the manuscript include the reverse genetics approaches to dissect the impact of specific male-killing loci, and use of a "masculinization" assay in Lepidopteran cell lines to determine the impact of interactions between specific masc and oscar homologs.

We would like to thank you for evaluating our manuscript.

Weaknesses:

It is clear from Figure 1 that the combinations of wmk homologs do not cause male killing on their own here. While I largely agree with the author's conclusions that oscar is the primary MK factor in this system, I don't think we can yet rule out that wmk(s) may work synergistically or interactively with oscar in vivo. This might be worth a small note in the discussion. (eg at line 294 'indicating that wmk likely targets factors other than masc." - this could be downstream of the impacts of oscar; perhaps dependent on oscar-mediated impacts on masc first).

We sincerely appreciate your suggestion. Whilst wmk genes themselves did not exhibit apparent lethal effects on the native host, as you noted, we cannot entirely rule out the possibility that wmk may be involved in male-killing actions, either directly or indirectly assisting the function of Hb-oscar. Following your suggestion, we have added a brief note in the discussion section regarding the interpretation of wmk functions.

“In addition, Katsuma et al. (2022) reported that the wmk homologs encoded by wFur did not affect the masculinizing function of masc in vitro, indicating that wmk likely targets factors other than masc. Whilst we cannot rule out the possibility that wmk may work synergistically or interactively with oscar in vivo—potentially acting downstream of oscar’s impact—our results strongly suggested that Wolbachia strains have acquired multiple MK genes through evolution.” (lines 287-292)

Regarding the perceived male-bias in Figure 2a: I think readers might be interpreting "unhatched" as "total before hatching". You could eliminate ambiguity by perhaps splitting the bars into male and female, and then within a bar, coloring by hatched versus unhatched. But this is a minor point, and I think the updated text helps clarify this.

Thank you for your suggestion. We have accordingly revised the figure 2a. In addition, we have included more detailed information in the first sentence of the section Males are killed mainly at the embryonic stage.

“The sex of hatched larvae (neonates) and the remaining unhatched embryos was determined by the presence or absence of W chromatin, a condensed structure of the female-specific W chromosome observed during interphase.” (lines 171-173)

The new Figure 4b looks to be largely redundant with the oscar information in Figure 1a.

Thank you for your suggestion. We have removed Figure 4b due to its overlap with Figure 1a and have incorporated relevant figure legends into the Figure 1a legend.

Updated statistical comparisons for the RNA-seq analysis are helpful. However these analyses are based on single libraries (albeit each a pool of many individuals), so this is still a weaker aspect of the manuscript.

Thank you for your suggestion. As you noted, the use of single libraries (due to the limited number of available individuals, though each includes approximately 50 males and females) may be a potential limitation of this study. However, as demonstrated in the qPCR assay for the Z-linked gene provided in the previous revision, we believe that our data and conclusion—that Wolbachia/ Hb-oscar disrupts dosage compensation by causing the overexpression of Z-linked genes—are well-supported and robust.

The new information on masc similarity is useful (Fig 4d) - if the authors could please include a heatmap legend for the colors, that would be helpful. Also, please avoid green and red in the same figure when key for interpretation.

Thank you for your suggestion. We have accordingly included a heatmap legend and revised the colors.

Figure 1A "helix-turn-helix" is misspelled. ("tern").

We have revised.

Recommendations for the authors:

Comments from the reviewing editor: I would suggest you address the comments of the reviewer on the revised version.

We have further revised the manuscript to address all the questions, comments and suggestions provided by the reviewers. We believe that the resulting revisions have significantly enhanced the quality and comprehensiveness of our manuscript.

Reviewer #1 (Recommendations for the authors):

Thank you for revising this manuscript. I have a few last recommendations:

- Line 214: re: 'Statistical data are available in the supplementary data file', it would be more helpful to add a few words here that actually summarize the statistical results

We would like to thank you for your suggestion. We have revised the sentence to describe the overview of the statistical results.

“RNA-seq analysis revealed that, in Hm-oscar-injected embryos, Z-linked genes (homologs on the B. mori chromosomes 1 and 15) were more expressed in males than in females (Fig. 3a), which was not observed in the GFP-injected group (Fig. 3b). Similarly, as previously reported by Arai et al. (2023a), high levels of Z-linked gene expression were also observed in wHm-t-infected males, but not in NSR males (Fig. 3c,d). The high (i.e., doubled) Z-linked gene expression in both Hm-oscar-expressed and wHm-t-infected males was further confirmed by quantification of the Z-linked Hmtpi gene (Fig. 3e). These trends were statistically supported, with all data available in the supplementary data file.” (lines 205-213)

- Figure 1 legend: do you mean 'bridged' instead of 'brigged'?

We have accordingly revise, thank you for the suggestion.

Changes from the earlier version to the revised version

Abstract

Abstract

Wolbachia are maternally transmitted bacterial symbionts that are ubiquitous among arthropods. They can hijack host reproduction in various ways, including male killing (MK), where the sons of infected mothers are killed during development. The recent discovery of MK-associated Wolbachia genes, i.e., oscar in Ostrinia moths and wmk in Drosophila flies, stimulates our interest in the diversity and commonality of MK mechanisms, which remain largely unclear. We recently discovered that a Wolbachia symbiont of the moth Homona magnanima carries an MK-associated prophage region encoding homologs of oscar (Hm-oscar) and wmk (wmk-1–4). Here, we investigated the effects of these genes in the native host. Upon transient overexpression, Hm-oscar, but not wmk, induced male lethality in H. magnanima, in contrast to our observations in Drosophila, where the wmk homologs, but not Hm-oscar, killed the males. Hm-oscar disrupted sex determination in male embryos by inducing a female-type doublesex splicing and impaired dosage compensation, recapitulating the Wolbachia phenotype. Cell-based transfection assays confirmed that Hm-oscar suppressed the function of masculinizer, the primary male sex determinant involved in lepidopteran dosage compensation. Our study highlights the conserved roles of oscar homologs in Wolbachia-induced lepidopteran MK and argues that Wolbachia have evolved multiple MK mechanisms in insects.

Introduction

Arthropods commonly carry microbial symbionts that are passed from mother to offspring (Hurst & Frost, 2015; Werren et al., 2008; Hurst, 2017). The maternally transmitted bacteria, belonging to the genus Wolbachia (Alphaproteobacteria), are estimated to be present in at least 40% of all insect species, making them one of the most widespread endosymbiont genera on the planet (Werren et al, 2008; Zug & Hammerstein, 2012). Wolbachia have achieved evolutionary success by manipulating host reproduction through various means that enhance endosymbiont transmission (Werren et al., 2008). Such manipulation of the host’s reproduction includes cytoplasmic incompatibility (CI), parthenogenesis, feminization, and male killing (MK), each of which affects the biological features, distribution, and evolution of the host. Among these strategies, MK directly skews the sex ratio of the host population toward females by killing male offspring of infected mothers during development. The lack of symbiont transmission through male hosts often leads to the evolution of MK. Wolbachia have been shown to induce MK in a wide range of insect taxa (e.g., dipterans, lepidopterans, and coleopterans) and other arthropods (Hurst et al., 1999; Kageyama & Traut, 2004; Hurst et al., 2000). Furthermore, various bacteria, viruses, and microsporidia induce MK phenotypes (Kageyama et al., 2012; Fujita et al., 2021; Kageyama et al., 2023; Nagamine et al., 2023), and recent studies have postulated that these microbes have evolved their MK ability independently (Harumoto & Lemaitre, 2018; Kageyama et al., 2023; Nagamine et al., 2023; Arai et al., 2023a).

Earlier version: The evolution and molecular mechanisms underlying Wolbachia-induced MK have attracted considerable attention for decades. It has been hypothesized that the mechanisms underlying MK are associated with sex determination cascades in insects (Hornett et al., 2022). Indeed, some MK-inducing Wolbachia strains have been shown to directly interact with the sex determination system in lepidopteran insects (Sugimoto and Ishikawa, 2012; Sugimoto et al., 2015; Fukui et al., 2015; Arai et al., 2023a). The Lepidopteran sex determination system consists of multiple transcriptional regulators, some of which exhibit sex-linked expression and/or splicing isoforms. Lepidoptera generally have a female heterogametic sex chromosome system (e.g., WZ/ZZ) and employ dosage compensation, which equalizes the sex-linked (Z-chromosome-linked) gene dose between males and females (Kiuchi et al., 2014). Dosage compensation is regulated by the masculinizer gene (masc), which is critical for male development. masc also controls the downstream master regulator of sex determination and differentiation, doublesex (dsx), which exhibits sex-dependent splicing isoforms (dsxF in females and dsxM in males) (Kiuchi et al., 2014). In Ostrinia and Homona male moths, MK-inducing Wolbachia strains impair dosage compensation by disrupting the expression of masc. Furthermore, they disrupt sex determination in male moths by inducing the “female” isoform of dsx (dsxF), leading to a “mismatch” between genetic sex (male: ZZ sex chromosome constitution) and phenotypic sex (female: based on dsxF) and consequently to male death (Arai et al., 2023a; Fukui et al., 2015; Sugimoto et al., 2015; Sugimoto & Ishikawa, 2012).

Revised version: The evolution and molecular mechanisms underlying Wolbachia-induced MK, observed across diverse insect taxa, have attracted considerable attention for decades. MK mechanisms are hypothesized to be linked to sex determination cascades regulating male and female differentiation in insects (Hornett et al., 2022). Indeed, some MK-inducing Wolbachia strains disrupt the sex determination system in host moths (Sugimoto and Ishikawa, 2012; Sugimoto et al., 2015; Fukui et al., 2015; Arai et al., 2023a). Lepidopteran insects (moths and butterflies) generally have a female heterogametic sex chromosome system (e.g., WZ in females, ZZ in males) and employ dosage compensation, which equalizes the Z-chromosome-linked gene dose between sexes (Kiuchi et al., 2014). Dosage compensation is regulated by the masculinizer gene (masc), which is critical for male development. The lepidopteran sex determination system consists of multiple transcriptional regulators, some of which exhibit sex-linked expression and/or splicing isoforms. The primary male sex determinant, masc, also regulates the downstream master sex determination gene, doublesex (dsx), which exhibits sex-dependent splicing isoforms (dsxF in females and dsxM in males) (Kiuchi et al., 2014). In Ostrinia and Homona male moths, MK-inducing Wolbachia strains impair dosage compensation by disrupting the expression of masc. Furthermore, they disrupt sex determination in male moths by inducing the “female” isoform of dsx (dsxF), leading to a “mismatch” between genetic sex (male: ZZ sex chromosome constitution) and phenotypic sex (female: based on dsxF), ultimately resulting in male death (Arai et al., 2023a; Fukui et al., 2015; Sugimoto et al., 2015; Sugimoto & Ishikawa, 2012).

Earlier version: More recently, a protein produced in Wolbachia wFur strain named Oscar (Osugoroshi protein containing CifB C-terminus-like domain and many Ankyrin Repeats; Osugoroshi translates to male killing in Japanese) was shown to recapitulate the Wolbachia-induced MK phenotype in the native host Ostrinia furnacalis (Katsuma et al., 2022). The Oscar protein interacts with and degrades Masc protein, leading to the failure of dosage compensation and the production of female-type dsx isoforms in Ostrinia male moths (Fukui et al., 2024; Katsuma et al., 2022). Although oscar homologs have been identified in several MK-inducing Wolbachia strains in Lepidoptera, some MK Wolbachia strains in Diptera (Drosophila) and Lepidoptera do not carry this gene (Arai et al., 2023b; Arai et al., 2024a). Furthermore, oscar homologs are evolutionarily dynamic, with highly variable sequences and structures (classified as type I to type III oscar based on sequence homology) (Arai et al., 2024a), making it difficult to assess their functional relevance. In addition to oscar, the helix-turn-helix domain-containing putative transcriptional regulator wmk is widely conserved among Wolbachia strains and induces various toxicities in Drosophila flies (i.e., no effect, weak male lethality [30% of males die], strong male lethality [90%–100% of males die], and death of all males and females) (Perlmutter et al., 2019; Perlmutter et al., 2020; Perlmutter et al., 2021; Arai et al., 2023b). Although the mechanisms underlying the wmk-induced toxicities and their connection to sex determination systems remain unclear, these findings suggest that Wolbachia strains carry multiple factors that cause male lethality. However, the diversity and commonality of these functions in insects remain largely unknown, partly due to the technical challenges in validating gene functions in non-model insects.

Revised version: More recently, a Wolbachia protein named Oscar (Osugoroshi protein containing CifB C-terminus-like domain and many Ankyrin Repeats; Osugoroshi translates to male killing in Japanese) was shown to recapitulate the Wolbachia wFur-induced MK phenotype in the native host moth Ostrinia furnacalis (Katsuma et al., 2022). The Oscar protein interacts with and degrades Masc protein, leading to the failure of dosage compensation and the production of female-type dsx isoforms in Ostrinia male moths (Fukui et al., 2024; Katsuma et al., 2022). Although oscar homologs are common in MK Wolbachia strains from Lepidoptera, MK strains from Diptera (Drosophila) and from several lepidopteran hosts lacked the gene (Arai et al., 2023b; Arai et al., 2024a). Furthermore, oscar homologs are evolutionarily dynamic, with highly variable sequences and structures (Arai et al., 2024a), making it difficult to assess their functional relevance. In addition to oscar, the helix-turn-helix domain-containing putative transcriptional regulator wmk is widely conserved among Wolbachia strains and induces various toxicities in Drosophila flies, ranging from no effect, weak male lethality (30% mortality), strong male lethality (90%–100% mortality), or complete lethality of both sexes (Perlmutter et al., 2019; Perlmutter et al., 2020; Perlmutter et al., 2021; Arai et al., 2023b). Although the mechanisms underlying the wmk-induced toxicities and their connection to sex determination systems remain unclear, these findings suggest that Wolbachia strains carry multiple factors that cause male lethality. However, the diversity and commonality of these functions in insects remain largely unknown, partly due to the technical challenges in validating gene functions in non-model insects.

We recently discovered a prophage region responsible for a Wolbachia-induced MK by comparing closely related non-MK Wolbachia (wHm-c) and MK Wolbachia (wHm-t) in the tea tortrix moth Homona magnanima (Tortricidae) (Arai et al., 2023b; Arai et al., 2024b). The MK-associated prophage element encodes potential MK causative genes–four wmk homologs (wmk-1 to wmk-4) as well as an oscar homolog (Hm-oscar, encoding 1181 aa protein), which differs significantly in sequence length and structure from the wFur-encoded oscar (encoding 1830 aa protein) (Fig. 1a). When these MK candidate genes are transgenically overexpressed in Drosophila, wmk-1 and wmk-3 have a lethality of almost 100%, while wmk-2, wmk-4, and Hm-oscar induce no lethal effects when singly overexpressed. Furthermore, co-expression of the adjacent wmk-3 and wmk-4 has been shown to induce the death of 90% of male flies and restores female survival, suggesting their combined action (Arai et al., 2023b). However, transgenic expression of genes from a moth-derived male killer in flies does not necessarily replicate the native host–bacteria interaction, and the mechanistic links between prophage-encoded Wolbachia genes and MK in the native host Homona remain unclear.

Transient expressions of the Hm-oscar gene resulted in female-biased sex ratios

Earlier version: (a) Structures of wFur-encoded Oscar and wHm-t-encoded Hm-Oscar and four Wmk proteins. (b) Male ratio of adult progeny obtained from cRNA-injected groups (wmk-1, wmk-2, T2A-brigged wmk-1 and wmk-2, wmk-3, wmk-4, T2A-brigged wmk-3 and wmk-4, GFP, and Hm-oscar; n=5–15 independent replicates using different egg masses), the non-injected NSR line, and the MK wHm-t-positive WT12 line. The total numbers of adult females and males are shown at the bottom. Different letters indicate significant differences (Steel–Dwass test, P < 0.05). The dot plots show all data points individually.

Revised version: (a) Structures of wFur-encoded Oscar and wHm-t-encoded Hm-Oscar and four Wmk proteins. (b) Male ratio of adult progeny obtained from cRNA-injected groups (wmk-1, wmk-2, T2A-bridged wmk-1 and wmk-2, wmk-3, wmk-4, T2A-bridged wmk-3 and wmk-4, GFP, and Hm-oscar; n = 5–15 independent replicates using different egg masses), the non-injected NSR line, and the MK wHm-t-positive WT12 line. The total numbers of adult females and males are shown at the bottom. Different letters indicate significant differences (Steel–Dwass test, P < 0.05). The dot plots show all data points individually.

In this study, we showed that the prophage-encoded Hm-oscar recapitulates Wolbachia-induced MK in H. magnanima. Furthermore, we clarified the mechanistic links to host sex-determination cascades both in vivo and in vitro and discussed the underlying mechanisms of MK in Lepidoptera, arguing for the diverse evolutionary origin of Wolbachia-induced MK.

Results

Hm-oscar induces female-biased sex ratios

To achieve the transient overexpression of Hm-oscar and the four wmk genes (wmk-1, wmk-2, wmk-3, and wmk-4), constructed mRNA (cRNA) was injected into Wolbachia-free H. magnanima embryos. Subsequently, the adult moths that emerged from the cRNA-injected embryos were sexed based on their external morphology. When Hm-oscar was overexpressed, the sex ratio of adults was strongly female-biased (42 males and 218 females in total, 14.4% ± 11.1% male ratio in 15 replicates), which was in sharp contrast with the ratios observed in the GFP-injected (126 females and 123 males in total, 49.1% ± 2.08% male ratio in 10 replicates) and non-injected (NSR) (201 females and 215 males in total, 51.4% ± 3.08% male ratio in 9 replicates) groups (P = 0.001 and 0.002, respectively, Steel–Dwass test, Fig. 1b). Compared with the GFP-injected group, the sex ratio was not biased by the overexpression of wmk-1 (51.1% ± 7.27% male ratio in 6 replicates, P = 0.999), wmk-2 (53.2% ± 4.77% male ratio in 6 replicates, P = 0.810), wmk-3 (53.3% ± 7.91% male ratio in 5 replicates, P = 0.999), and wmk-4 (52.9% ± 3.47% male ratio in 5 replicates, P = 0.773). Although the dual expression of wmk-3 and wmk-4 induces strong male lethality in Drosophila (Arai et al., 2023b), here, it did not skew the sex ratio of Homona compared with the ratio detected in the GFP-injected group (57.9% ± 7.67% male ratio in 5 replicates, P = 0.704). Similarly, the dual expression of the tandemly arrayed wmk-1 and wmk-2 did not bias the sex ratio (49.4% ± 2.82% male ratio in 5 replicates) compared with the value in the GFP-injected group (P = 1.000). The sex ratio data of all replicates are available on the supplementary data file.

Males are killed mainly at the embryonic stage

Earlier version: The sex of unhatched embryos and hatched larvae (neonates) was determined by karyotyping for W chromatin, with the presence and absence of this substance indicating females and males, respectively. MK Wolbachia wHm-t kills Homona males during embryogenesis, resulting in a female-biased sex ratio in hatched larvae (neonates) and a male-biased sex ratio in unhatched embryos (Arai et al., 2023b; Arai et al., 2020). In line with this finding, the sex ratios of hatched larvae (34 females and 6 males) and unhatched embryos (27 females and 47 males) of the MK wHm-t infected line (WT12) were also female- and male-biased, respectively (P = 0.013 and 0, respectively). In the Hm-oscar-injected group, the sex ratio of the hatched larvae (neonates) was strongly female-biased (21 females and 3 males, P = 0.0001 in the binomial test) (Fig. 2a). In contrast, the unhatched mature embryos showed a male-biased ratio (18 females and 38 males, P = 0.005), suggesting that male mortality primarily occurred during the embryonic stage, reflecting the wHm-t-induced MK phenotype. In the non-injected group (NSR), the sex ratios of hatched larvae (28 females and 28 males) and unhatched embryos (23 females and 25 males) were not biased (P = 0.553 and 0.443, respectively).

Revised version: The sex of hatched larvae (neonates) and the remaining unhatched embryos was determined by the presence or absence of W chromatin, a condensed structure of the female-specific W chromosome observed during interphase. MK Wolbachia wHm-t kills Homona males during embryogenesis, resulting in a female-biased sex ratio in hatched larvae (neonates) and a male-biased sex ratio in unhatched embryos (Arai et al., 2023b; Arai et al., 2020). In line with this finding, the sex ratios of hatched larvae (34 females and 6 males) and unhatched embryos (27 females and 47 males) of the MK wHm-t infected line (WT12) were also female- and male-biased, respectively (P = 0.013 and 0, respectively). In the Hm-oscar-injected group, the sex ratio of the hatched larvae (neonates) was strongly female-biased (21 females and 3 males, P = 0.0001 in the binomial test) (Fig. 2a). In contrast, the unhatched mature embryos showed a male-biased ratio (18 females and 38 males, P = 0.005), suggesting that male mortality primarily occurred during the embryonic stage, reflecting the wHm-t-induced MK phenotype. In the non-injected group (NSR), the sex ratios of hatched larvae (28 females and 28 males) and unhatched embryos (23 females and 25 males) were not biased (P = 0.553 and 0.443, respectively).

Hm-Oscar induced lethality in male embryos with the female-type sex determination

(a) Sex ratio of the hatched larvae and unhatched embryos in the Hm-oscar-expressed, wHm-t-infected, and non-infected/expressed groups. Females and males were discriminated based on the presence or absence of W chromatin. The number of individuals is indicated in brackets. (b) Splicing patterns of the downstream sex-determining gene dsx of H. magnanima embryos (5 days post oviposition). Abbreviations: HmOsc, Hm-oscar injected group; GFP, GFP-injected group; WT12, wHm-t-infected line; NSR, non-infected/injected line. M and F indicate W chromatin-negative (ZZ: male genotype) and W chromatin-positive (ZW: female genotype) mature embryos, respectively. dsx-F and dsx-M represent female and male-specific splicing variants, respectively. L: 100 bp DNA ladder (ExcelBand 100 bp DNA Ladder, SMOBIO Technology, Inc., Hsinchu, Taiwan). 0.1, 0.5, 1.0, and 1.5 kb markers are indicated with arrows. N: negative control (water).

Female-type sex determination in male embryos that are destined to die

Splicing patterns of the downstream sex determinant dsx were assessed in each of the H. magnanima embryos that were Hm-oscar-expressed, GFP-expressed, wHm-t-infected, or non-expressed (i.e., non-injected, NSR). In GFP-expressed and NSR embryos, females and males (which were identified by the presence and absence of W chromatin, respectively) exhibited a female- and male-type splicing variant of dsx, respectively. However, both wHm-t-infected and Hm-oscar-expressed embryos induced female-type dsx splicing regardless of the presence or absence of W chromatin (Fig. 2b). Notably, male embryos expressing Hm-oscar also exhibited weak male-type dsx splicing in addition to the female-type splicing, resembling the previously observed pattern in male embryos infected with low-titer wHm-t (Arai et al., 2023a).

Hm-oscar impairs dosage compensation in male embryos

Earlier version: To confirm the effects of Hm-oscar on dosage compensation, we quantified the expression of Z-linked genes using RNA-seq and qPCR. Unlike other lepidopteran species, the family Tortricidae, including H. magnanima, generally possess a large Z chromosome that is homologous to B. mori chromosomes 1 (Z) and 15 (autosome). RNA-seq analysis revealed that, in Hm-oscar-injected embryos, Z-linked genes (homologs on the B. mori chromosomes 1 and 15) were more expressed in males than in females (Fig. 3a), which was not observed in the GFP-injected group (Fig. 3b). Similarly, as previously reported by Arai et al. (2023a), high levels of Z-linked gene expression were also observed in wHm-t-infected males, but not in NSR males (Fig. 3c,d). Statistical data are available on the supplementary data file. The high (i.e., doubled) Z-linked gene expression in both Hm-oscar-expressed and wHm-t-infected males was further confirmed by quantification of the Z-linked Hmtpi gene (Fig. 3e).

Revised version: To confirm the effects of Hm-oscar on dosage compensation, we quantified the expression of Z-linked genes using RNA-seq and qPCR. Unlike other lepidopteran species, the family Tortricidae, including H. magnanima, generally possess a large Z chromosome that is homologous to B. mori chromosomes 1 (Z) and 15 (autosome). RNA-seq analysis revealed that, in Hm-oscar-injected embryos, Z-linked genes (homologs on the B. mori chromosomes 1 and 15) were more expressed in males than in females (Fig. 3a), which was not observed in the GFP-injected group (Fig. 3b). Similarly, as previously reported by Arai et al. (2023a), high levels of Z-linked gene expression were also observed in wHm-t-infected males, but not in NSR males (Fig. 3c,d). The high (i.e., doubled) Z-linked gene expression in both Hm-oscar-expressed and wHm-t-infected males was further confirmed by quantification of the Z-linked Hmtpi gene (Fig. 3e). These trends were statistically supported, with all data available in the supplementary data file.

Hm-oscar-overexpressed male embryos showed higher levels of Z-linked gene expression

(a-d) Normalized expression levels (TPM) and chromosomal distributions of transcripts in H. magnanima embryos. RNA-seq data of embryos (108 hpo) were used to make the following comparisons: Hm-oscar-injected males versus Hm-oscar injected females (a), GFP-injected males versus GFP-injected females (b), WT12 males versus WT12 females (c), and NSR males versus NSR females (d). The chromosome number for each H. magnanima transcript-derived contig was assigned based on Bombyx mori gene models. The X axis represents the chromosome number of B. mori (shown as chr01 to chr28), and chr01 and chr15 (highlighted in orange) correspond to the Z chromosome of H. magnanima (Arai et al., 2023a). (e) Relative expression of a Z-linked tpi gene in H. magnanima groups. Statistical significance is highlighted with different alphabets.

Hm-oscar suppresses the masculinizing functions of lepidopteran masc genes

Earlier version: To confirm whether Hm-oscar suppresses the functions of the primary male sex determinant masc, their interactions were assessed by transfection using BmN-4 cells, which are derived from B. mori ovaries (Grace, 1967) and exhibit a female-type default sex determination (Kiuchi et al., 2014). Previous studies have shown that transfection of BmN-4 cells with masc cDNA from various lepidopteran insects, such as O. furnacalis, induces the expression of BmdsxM and BmIMPM (Fukui et al., 2015; Katsuma et al., 2019; Katsuma et al., 2022). Since BmIMPM regulates the splicing of BmdsxM and is exclusively expressed in males, we selected BmIMPM as a quantitative marker of masculinization, as described in Katsuma et al. (2022). In contrast with the control groups, which were transfected with the non-inserted pIZ/V5-His vector and showed this default sex determination, the cells transfected with the Hmmasc-inserted pIZ/V5-His vector exhibited the male-type sex determination, as indicated by the increased expression of the male-specific splicing variant of the sex-determining gene BmImpM, although this increase was not statistically significant (Fig. 4a). Furthermore, the masculinizing function of Hmmasc was suppressed by co-transfecting the Hmmasc/Hm-oscar-inserted pIZ/V5-His vectors, as manifested by the low expression levels of BmImpM. wFur-encoded Oscar, which differs from Hm-Oscar in structure and amino acid sequence (Fig. 4b), also suppressed Hmmasc, the function of this gene was less active than that of Hm-oscar (Fig. 4a). Hm-oscar also suppressed the masculinizing function of masc genes derived from various lepidopteran insects (i.e., O. furnacalis, Spodoptera frugiperda, Bombyx mori, Trichoplusia ni, and Papilio machaon), which are highly divergent in their amino acid sequences (Fig. 4c,d). These results suggest that Hm-oscar has a broad spectrum of activity in Lepidoptera.

Revised version: To confirm whether Hm-oscar suppresses the functions of the primary male sex determinant masc, their interactions were assessed through transfection in BmN-4 cells. These cells, derived from B. mori ovaries (Grace, 1967), exhibit female-type default sex determination (Kiuchi et al., 2014). However, overexpression of the masc gene from various lepidopteran insects, such as O. furnacalis, induces male-type sex determination by promoting the expression of BmdsxM and BmIMPM in BmN-4 cells (Fukui et al., 2015; Katsuma et al., 2019; Katsuma et al., 2022). Since BmIMPM regulates the splicing of BmdsxM and is exclusively expressed in males, we selected BmIMPM as a quantitative marker of masculinization, as described in Katsuma et al. (2022).

Added in the revised version: In contrast with the non-expressed control groups, which exhibited default female-type sex determination, cells expressing Hmmasc exhibited male-type sex determination (Fig. 4a). This was indicated by the increased expression of the male-specific splicing variant BmIMPM, although this increase was not statistically significant. Furthermore, the masculinizing function of Hmmasc was suppressed when Hmmasc and Hm-oscar were co-expressed, as evidenced by the low expression levels of BmImpM (Fig. 4a). wFur-encoded Oscar, which differs from Hm-Oscar in structure and amino acid sequence (Fig. 1a), also suppressed Hmmasc, but its function was less active than that of Hm-oscar (Fig. 4a). Hm-oscar also suppressed the function of masc genes derived from various lepidopteran insects (i.e., O. furnacalis, Spodoptera frugiperda, Bombyx mori, Trilocha varians, and Papilio machaon), which are highly divergent in their amino acid sequences (Fig. 4b,c). These results suggest that Hm-oscar has a broad spectrum of activity in Lepidoptera.

Hm-oscar suppressed the masculinizing function of lepidopteran Masc

Earlier version: (a) Relative expression levels of the male-specific BmImpM variant in transfected BmN-4 cells. The relative expression of BmImpM in the control (non-inserted pIZ/V5) and Hm-oscar/oscar co-transfected groups was normalized by setting the mean in the Hmmasc single transfected group as 100. (b) Protein structures and homologies of Hm-Oscar and Oscar proteins. (c) Relative expression levels of the BmImpM in Hm-oscar/masc transfected BmN-4 cells. The relative expression of BmImpM in each condition was normalized by setting the mean in the Hmmasc single transfected group as 100. The dot plots show all data points individually. Each experimental condition was replicated three times. Abbreviations: Hm, Homona magnanima; Bm, Bombyx mori; Of, Ostrinia furnacalis; Pm, Papilio machaon; Sf, Spodoptera frugiperda; and Tv, Trilocha varians. (d) Phylogeny and homologies of the Masc protein sequences used in this study.

Revised version: (a) Relative expression levels of the male-specific BmImpM variant in transfected BmN-4 cells. The relative expression of BmImpM in the control (non-inserted pIZ/V5) and Hm-oscar/oscar co-transfected groups was normalized by setting the mean in the Hmmasc single transfected group as 100. (b) Relative expression levels of the BmImpM in Hm-oscar/masc transfected BmN-4 cells. The relative expression of BmImpM in each condition was normalized by setting the mean in the Hmmasc single transfected group as 100. The dot plots show all data points individually. Each experimental condition was replicated three times. Abbreviations: Hm, Homona magnanima; Bm, Bombyx mori; Of, Ostrinia furnacalis; Pm, Papilio machaon; Sf, Spodoptera frugiperda; and Tv, Trilocha varians. (c) Phylogeny and homologies of the Masc protein sequences used in this study. The sequence homology between each pair is presented in the figure and visualized as a heatmap, with red indicating high homology (Max = 100) and blue indicating low homology (Min = 23.82).

Discussion

In this study, we showed that the phage-encoded Hm-oscar, but not wmk, induced male lethality and a female-biased sex ratio in H. magnanima. Furthermore, the overexpressed Hm-oscar impaired male sex determination in Homona, recapitulating the wHm-t-induced phenotypes. Cell-based assays confirmed that Hm-oscar suppressed the masculinizing functions of masc. These results strongly suggested that Hm-oscar underlies the wHm-t-induced MK function in H. magnanima. In light of our results and previous findings in Homona, Ostrinia, and Bombyx (Arai et al., 2023a; Katsuma et al., 2022; Kiuchi et al., 2014), we hypothesize the following molecular mechanisms for wHm-t-induced MK: The Hm-Oscar protein targets and inhibits the function of the HmMasc protein, thereby resulting in female sex determination and disrupting dosage compensation, which ultimately leads to male death. Although the full mechanisms remain unclear, unbalanced Z-linked gene expression mediated by improper dosage compensation—resulting in Z-linked gene expression levels twice as high as those in normal males—appear to be lethal for lepidopteran males (Kiuchi et al., 2014; Fukui et al., 2015; Visser et al., 2021).

The Wolbachia-encoded oscar homologs identified so far show diversity in sequence length and are classified as type I, II, and III based on their amino acid sequence homologies. Type I and II Oscar proteins, derived from MK-inducing Wolbachia, contain many ankyrin repeats and a proteinase domain (Arai et al., 2024a). Although long ankyrin repeat sequences are postulated to be critical for the function of wFur-encoded Oscar protein (1830 aa, type I) (Katsuma et al., 2022), our study revealed that the functions of Hm-Oscar protein (1181 aa, type II), which encodes fewer ankyrin repeats, were comparable with those of Oscar protein carried by wFur. Interestingly, type II oscar gene is also present in the feminizing Wolbachia wFem in Eurema butterflies (Arai et al., 2024a). oscar homologs, which inhibit the masculinizing function and induce female sex determination, may have a conserved function in Wolbachia-induced MK and feminization in Lepidoptera.

Earlier version: In contrast to the results of this study, we have previously demonstrated that the phage-encoded wmk, but not Hm-oscar, induces male lethality in D. melanogaster (Arai et al., 2023b). Although the means of expression are different (i.e., transient in Homona and transgenic in Drosophila), this finding highlighted the differences in the mode/range of action of Wolbachia genes between insect species. It has been hypothesized that microbes induce MK in insects by targeting molecular mechanisms involved in sex determination and differentiation (Hornett et al., 2022). Sex determination systems in insects are diverse; for example, Lepidoptera (including H. magnanima) and Diptera (including D. melanogaster) do not share any known sex-determining genes other than dsx (Suzuki, 2018). The different outcomes in Homona and Drosophila are probably due to their different sex determination systems. Because oscar interacts with and suppresses masc, Hm-oscar could induce mortality in Homona males that possess masc, but not in Drosophila males that lack it. Considering that oscar homologs are not present in known MK Wolbachia strains in dipteran insects (Arai et al., 2024c; Katsuma et al., 2022), the mechanisms (i.e., causative genes) of Wolbachia-induced MK probably differ between insect taxa (e.g., between Lepidoptera and Diptera). While the mechanisms underlying wmk-induced lethality remain unclear, the distinct effects associated with this gene between Homona and Drosophila may also reflect their genetic backgrounds (e.g., presence/absence of host factor(s) that interact with wmk). In addition, Katsuma et al. (2022) reported that the wmk homologs encoded by wFur did not affect the masculinizing function of masc in vitro, indicating that wmk likely targets factors other than masc. Our results strongly suggested that Wolbachia strains have acquired multiple MK genes through evolution.

Revised version: In contrast to the results of this study, we have previously demonstrated that the phage-encoded wmk, but not Hm-oscar, induces male lethality in D. melanogaster (Arai et al., 2023b). Although the means of expression are different (i.e., transient in Homona and transgenic in Drosophila), this finding highlighted the differences in the mode/range of action of Wolbachia genes between insect species. It has been hypothesized that microbes induce MK in insects by targeting molecular mechanisms involved in sex determination and differentiation (Hornett et al., 2022). Sex determination systems in insects are diverse; for example, Lepidoptera (including H. magnanima) and Diptera (including D. melanogaster) do not share any known sex-determining genes other than dsx (Suzuki, 2018). The different outcomes in Homona and Drosophila are probably due to their different sex determination systems. Because oscar interacts with and suppresses masc, Hm-oscar could induce mortality in Homona males that possess masc, but not in Drosophila males that lack it. Considering that oscar homologs are not present in known MK Wolbachia strains in dipteran insects (Arai et al., 2024c; Katsuma et al., 2022), the mechanisms (i.e., causative genes) of Wolbachia-induced MK probably differ between insect taxa (e.g., between Lepidoptera and Diptera). While the mechanisms underlying wmk-induced lethality remain unclear, the distinct effects associated with this gene between Homona and Drosophila may also reflect their genetic backgrounds, such as the presence or absence of host factor(s) that interact with wmk. In addition, Katsuma et al. (2022) reported that the wmk homologs encoded by wFur did not affect the masculinizing function of masc in vitro, indicating that wmk likely targets factors other than masc. Whilst we cannot rule out the possibility that wmk may work synergistically or interactively with oscar in vivo— potentially acting downstream of oscar’s impact—our results strongly suggested that Wolbachia strains have acquired multiple MK genes through evolution.

Wolbachia-induced phenotypes are known to be influenced by the genetic backgrounds of hosts (Hornett et al., 2006; Sasaki et al., 2002; Veneti et al., 2012). Our study showed that the wHm-t-encoded Hm-oscar suppresses the function of masc in H. magnanima more efficiently than the wFur-encoded oscar, suggesting that this Wolbachia factor has undergone evolutionary tuning to adapt to its natural host. However, the mere presence of oscar and wmk homologs does not ensure MK expression. For example, wCauA, bearing type I oscar, did not induce MK in Ephestia (Cadra) cautella collected around 2000, although it did induce MK when transferred to the closely related host Ephestia kuehniella (Sasaki et al., 2002). An MK phenotype, presumably induced by wCauA, was also recorded in E. cautella around the 1970s (Takahashi & Kuwahara, 1970). These findings suggest a suppressor evolution against wCauA-induced MK in E. cautella. In addition, Wolbachia typically encode a number of wmk homologs in their genomes, but their effects differ among insects and are not apparent in some species such as Homona, possibly due to the diversity in sex determination system or the evolution of suppressors. Nonetheless, given the high copy number of wmk homologs in Wolbachia genomes, they may still contribute to Wolbachia biology as transcriptional regulators and potential virulence factors. Virulence genes often undergo duplication and substitution under strong selective pressure (Hill et al., 2022; Jones & Dangl, 2006). An intense evolutionary arms race between Wolbachia and their hosts could have increased and diversified the MK-associated Wolbachia genes. Conversely, natural selection favors the rescue of males by suppressing the Wolbachia-induced reproductive manipulations (Hornett et al., 2006; Hornett et al., 2014), which may involve changes in the sex determination system because Wolbachia strains frequently hijack host reproduction systems. In this context, Wolbachia and other MK-inducing microbes may be a hidden driver for the diversification of complex insect sex determination systems.

In this study, we clarified the conserved roles of the Wolbachia-encoded oscar homologs in Lepidoptera and demonstrated that Wolbachia have evolved distinct MK mechanisms (through causative genes) in insect taxa. The diversification of phenotype/virulence-associated genes and the rampant horizontal transmission of phages carrying virulence genes between Wolbachia strains may have contributed to the outstanding success of this bacterial genus. In addition to oscar and wmk, Wolbachia may retain other uncharacterized genes that induce male lethality, and further studies on diverse Wolbachia–host systems are highly warranted. Our findings provide insights into the molecular mechanisms and evolutionary relationships between endosymbionts and their hosts, which may also contribute to the design of pest management strategies.

Materials and methods

Experimental model and subject details

Homona magnanima

A laboratory-maintained H. magnanima line with a normal sex ratio (NSR) that was negative for Wolbachia and other endosymbionts, was used in our experiments. This line was initially collected in Tokyo, Japan, in 1999 and has been maintained inbred for over 250 generations in the laboratory. Larvae were reared using the artificial SilkMate 2S diet (Nosan Co., Yokohama, Japan) at 25°C under a long photoperiod (16L:8D) (Arai et al., 2019).

The laboratory-maintained all female H. magnanima WT12 line, which was initially collected in Taiwan (Tea Research Extension Station, Taoyuan city) in 2015 (Arai et al., 2020), was also used in this study. This line was maintained for over 50 generations by crossing it with the males of the NSR line.

BmN-4 cell line

BmN-4 cells (provided by Chisa Yasunaga-Aoki, Kyushu University, and maintained in our laboratory) were cultured at 26°C in IPL-41 Insect Medium (Applichem, Darmstadt, Germany) supplemented with 10% fetal bovine serum.

Transient expression of MK-associated phage genes

mRNA synthesis

Codon-optimized wmk genes (wmk-1 to wmk-4), conjugated wmk pairs using a T2A peptide (i.e., wmk-1-T2A-wmk-2 and wmk-3-T2A-wmk-4), and Hm-oscar genes synthesized by Arai et al. (2023b) were used for mRNA synthesis. These synthetic genes were ligated into the plasmid pIZ/V5-His (Invitrogen, MA, USA) using the NEBuilder® HiFi DNA Assembly kit (New England Biolabs, MA, USA) following the manufacturer’s protocol. The inserts of the vector (i.e., wmk-1, wmk-2, wmk-3, wmk-4, Hm-oscar, wmk-1-T2A-wmk-2, wmk-3-T2A-wmk-4, and GFP as a control gene) were amplified using KOD-one (TOYOBO, Osaka, Japan) with the primer set containing the T7 promoter described in Fukui et al. (2015) (i.e., pIZ-F-T7: 5′-TAATACGACTCACTATAGGGAGACAGTTGAACAGCATCTGTTC-3′ and pIZ-R: 5′-GACAATACAAACTAAGATTTAGTCAG-3′) under the following PCR conditions: 20 cycles at 98°C for 10 s, 62°C for 5 s, and 68°C for 15 s. The amplicons were purified using the QIAquick PCR purification kit (Qiagen, Hilden, Germany), and 100–200 ng DNA was used for mRNA synthesis. The capped mRNA (cRNA) with poly(A) tail was synthesized using the mMESSAGE mMACHINE® T7 Ultra kit (Invitrogen, MA, USA) following the manufacturer’s protocol with some modifications. In brief, the assembled transcription reaction (10 µL of T7 2X NTP/ARCA, 2 µL of 10X T7 reaction buffer, 2 µL of T7 enzyme mix, and 6 µL of linear DNA template diluted with nuclease-free water) was incubated at 37°C for at least 10 h to maximize RNA yields. After poly(A) tailing, cRNA was purified using ISOGEN II (Nippongene, Tokyo, Japan) and dissolved in up to 10 µL of nuclease-free water to achieve an RNA concentration of approximately 1500–2000 ng/µL. The synthesized cRNA was preserved at −80°C until further use.

Preparation of H. magnanima eggs

A total of 15 males and 10 females of H. magnanima were mated in a plastic box (30 cm × 20 cm × 5 cm) for 3–4 days. The collection of egg masses began the day after oviposition was confirmed. In brief, newly oviposited egg masses were collected at 30-min intervals during the dark period using red light (which the moths cannot perceive). Females started to oviposit eggs at least 4 h into the dark period and, within less than 30 min post oviposition (mpo), the egg masses were collected. The collection lasted until the start of the light period. The egg masses were then subjected to microinjection assays.

Inoculum preparation, microinjection, and maintenance of the injected embryos

A glass needle for microinjection was prepared using glass capillary GD-1 (Narishige, Tokyo, Japan) with a PC-10 puller (Narishige). The glass capillaries were pulled at two temperatures (first stage: 75°C, second stage: 65°C) using two heavy weights and one light weight. The movement position during the second heating stage was set to 3 mm (range 1–10 mm).

The cRNA solution was diluted in a buffer (100 mM potassium acetate, 2 mM magnesium acetate, and 30 mM HEPES-KOH; at pH 7.4) containing 0.2% (W/V) Brilliant Blue FCF (Wako, Osaka, Japan) to obtain an RNA concentration of 1000 ng/µL. Approximately 1–4 µL of dye-containing mRNA solution was aspirated into the glass needles (capillaries), and the edge of each needle was ground using Micro Grinder EG-402 (Narishige) at an angle of 20 degrees for 2 s.

The fresh egg masses of H. magnanima (collected at less than 30 mpo as described above) were put on a double-sided sticky tape (15 mm × 5 m T4612, Nitoms, Tokyo, Japan) and placed on glass slides. Under a Nikon SMZ1270 microscope (Nicon, Tokyo, Japan), the RNA solution (30–100 fL) was injected into individual eggs using an microinjector IM-400 (Narishige) (balance pressure set at 0.010–0.030; injection pressure set at 0.050-0.100) and a QP-3JOY-2R electric micromanipulator (MicroSupport, Shizuoka, Shizuoka).

The injected egg masses were maintained in a 90-mm plastic Petri dish fitted with a slightly wet filter paper. As high humidity interfered with the development of the embryos, the injected eggs (egg mass) were first maintained at a relative humidity of 30%–50% (0–3 days post injection) and then at a higher humidity until hatching (4–6 dpo, 60% relative humidity). The hatched larvae were reared separately with 1/2 ounces of SilkMate 2S (Nosan Co.) until eclosion. The adult H. magnanima moths that emerged from the cRNA-injected embryos were sexed based on their external morphology.

Sexing of embryos/neonates and RNA extraction

To verify the effect of the transient expression of Hm-oscar on sex determination in H. magnanima, Hm-oscar/GFP-expressed, non-injected, and wHm-t-infected mature embryos showing black head capsule (1 day before hatching, 5–6 days post oviposition, dpo) were dissected on glass slides using forceps, as described in Arai et al. (2022). Their Malpighian tubules were fixed with methanol/acetic acid (50% v/v) and stained with lactic acetic orcein for W chromatin observations. The remaining tissues not used for sexing were stored in ISOGEN II (Nippon Gene) at −80°C until subsequent extraction. In total, 12 males or females (confirmed based on the presence or absence of W chromatin) were pooled and homogenized in ISOGEN II to extract RNA as described in Arai et al. (2023a). In brief, 0.4 times the volume of UltraPure distilled water (Invitrogen) was added to the ISOGEN II homogenates, which were centrifuged at 12000 × g for 15 min at 4°C to pellet proteins and DNAs. The resulting supernatant was mixed with the same volume of isopropanol to precipitate RNAs; then, the resulting solutions were transferred to EconoSpin columns (Epoch Life Science) and centrifuged at 17,900 × g and 4°C for 2 min. The RNAs captured in the column were washed twice with 80% ethanol and eluted in 20 µL of UltraPure distilled water (Invitrogen).

Hmdsx detection

Sex-specific dsx splicing variants of H. magnanima were detected as described in Arai et al. (2023a). In brief, total RNA (100–300 ng) extracted from sex-determined mature embryos was reverse-transcribed using PrimeScript− II Reverse Transcriptase (TaKaRa, Shiga, Japan) at 30°C for 10 min, 45°C for 60 min, and 70°C for 15 min. Then, cDNA was amplified using KOD-FX Neo (Toyobo Co., Ltd.) with the following two primers: Hmdsx_long3F (5′-TGCCTAAAGTGAAAACGCCGAGGAGCC-3′) and Hmdsx_Mrev (5′-TGGAGGTCTCTTTTCATCCGG-3′). The PCR conditions used were as follows: 94°C for 2 min, followed by 45 cycles of 98°C for 10 s, 66°C for 30 s, and 68°C for 30 s. The amplicons were subjected to electrophoresis on 2.0% agarose Tris-borate-EDTA buffer (89 mM Tris-borate, 89 mM boric acid, 2 mM EDTA) gels.

Quantification of Z chromosome-linked genes

The effects on dosage compensation were assessed by measuring differences in gene expression, as described in Fukui et al. (2015) and Arai et al. (2023a). A total of 1.0 µg of the total RNA extracted from Hm-oscar or GFP-overexpressed mature embryos (108 hpo, approximately 40-50 sex determined embryos) was used to prepare mRNA-seq libraries via the NEBNext Poly (A) mRNA Magnetic Isolation Module (New England Biolabs) and the NEBNext Ultra II RNA Library Prep kit for Illumina (New England Biolabs) following the manufacturer’s protocol. The adaptor sequences and low-quality reads (Qscore <20) were removed from the generated sequence data [150 bp paired-end (PE150)] using Trimmomatic (Bolger et al., 2014). The trimmed reads were mapped to the previously assembled transcriptome database for H. magnanima (Arai et al., 2023a) using Kallisto (Bray et al., 2016) to generate the normalized read count data (transcripts per million, TPM). While we did not generate new RNA-seq data of the wHm-t-infected and uninfected groups, the expression data included in Arai et al. (2023a) were re-analyzed along with the newly obtained RNA-seq data above. The binary logarithms of the TPM differences between males and females belonging to each H. magnanima group (i.e., Hm-oscar/GFP-expressed) were calculated to assess the fold-changes in gene expression levels between sexes. As described in Arai et al. (2023a), the transcriptome data of H. magnanima were annotated using the B. mori gene sets obtained from KAIKOBASE (https://kaikobase.dna.affrc.go.jp). The binary logarithms of the TPM differences between males and females in the B. mori chromosomes 1 to 28 were plotted.

To confirm the effects observed in the RNA-seq-based quantification of the Z chromosome genes, expression level of the Z-linked Hmtpi gene was quantified by reverse transcription-qPCR as described in Arai et al. (2023a). Briefly, total RNA (100– 300 ng) extracted from pooled male or female mature embryos (108 hpo) was reverse transcribed using PrimeScript II reverse transcriptase (TaKaRa) as described above. The cDNA was used to quantify relative gene expression levels, with normalization to the control gene elongation factor 1a (ef1a). qPCR was performed using primer sets for Hmtpi (HmTpi_F180702_212: 5′-GCTGCGAGTGGGTGATTTTG-3′; HmTpi_R180702_326: 5′-GCGATCACTTTCAGACCCGA-3′) and Hmef1a (Hmef1a_F_val1_85: 5′-TTTCCAGGGTGGTTGAGCA-3′; Hmef1a_R_val1_193: 5′-CCGTTAAGGAGCTGCGTCG-3′), and KOD SYBR qPCR Mix (Toyobo, Osaka, Japan) in a LightCycler 96 system (Roche, Basel, Switzerland). The qPCR reaction was made with 5 µl KOD SYBR (TOYOBO), 0.4 µL forward primer [10 pmol/µl], 0.4 µl reverse primer [10 pmol/µl], and 3.2 µl water. The qPCR was performed under the following conditions: 180 s at 95°C, 40 cycles of 8 s at 98°C, 10 s at 60°C, and 10 s at 68°C, followed by heating to 90°C for melting curve analysis. Mean cycle threshold (CT) values of samples were calculated for at least 13 replicates, and both ΔCT (CT Z gene – CT ef1a) and ΔΔCT (ΔCTmale – ΔCTfem) values were calculated.

Transfection assays and quantification of BmIMPM

The coding sequence of Hmmasc was amplified from cDNA derived from the RNA extracted from male embryos of the NSR line using KOD-one (TOYOBO) with the following primer set: HmMasc_CDS_HindIIIF: 5′-GCAAAGCTTCAACATGATCTCTCGCCAACCACAATCAACATCA-3′ and HmMasc_CDS_BamHI R: 5′-GCAGGATCCCAACCTACTGATAAGGAGGGAAGTAAGGCTGCTG-3′. The following PCR conditions were applied: 20 cycles of 98°C for 10 s, 62°C for 5 s, and 68°C for 15 s. The amplicon was purified with the QIAquick PCR purification kit (Qiagen) and cloned into the pIZ/V5-His vector using the HindIII and BamHI enzymes (New England Biolabs). The codon-optimized oscar gene of wFur (Katsuma et al. 2022) was also cloned into the pIZ/V5-His vector using the KpnI and NotI enzymes (TaKaRa).

To verify the masculinizing function of Hmmasc, BmN-4 cells (4 × 105 cells per dish, diameter 35 mm) were transfected with 1 µg of each plasmid DNA (pIZ/V5-His having Hmmasc) using FuGENE HD (Promega, WI, USA), as described in Katsuma et al. (2022). To clarify whether Hm-oscar suppressed the function of Hmmasc, 1 µg of plasmid DNA (pIZ/V5-His bearing either Hmmasc or Hm-oscar and non-inserted vector as negative control) was transfected to the BmN-4 cells using FuGENE HD (Promega). Three days after transfection, the cells were collected and subjected to RNA extraction via TRI REAGENT® (Molecular Research Center Inc., USA) and cDNA construction with AMV transcriptase (TaKaRa). The degree of masculinization in the BmN-4 cells (default female-type sex determination) was verified by quantifying the expression levels of BmIMPM (Suzuki et al., 2010), which is involved in male-specific sex determination cascades, using primers BmIMP_F: 5′-ATGCGGGAAGAAGGTTTTATG-3′ and BmIMP_R: 5′-TCATCCCGCCTCAGACGATTG-3′, as described in Fukui et al. (2015). BmRp49 gene was also amplified as a control gene using primers rp49_F: 5′-CCCAACATTGGTTACGGTTC-3′ and rp49_R: 5′-GCTCTTTCCACGATCAGCTT-3′. Further interactions between Hm-oscar and masc genes derived from lepidopteran insects [i.e., Trilocha varians masc (Tvmasc), Spodoptera frugiperda masc (Sfmasc), B. mori masc (Bmmasc), O. furnacalis masc (Ofmasc), and Papilio machaon masc (Pmmasc)] (Katsuma et al., 2022b) were assessed using the same procedures. The expression levels of BmIMPM were analyzed as described in Katsuma et al. (2022). In brief, the relative expression of BmIMPM (CT BmIMPM / CT BmRp49) was calculated for each experimental group. Then, relative expression of BmIMPM of each sample was normalized by setting the mean in the Hmmasc singly transfected group as 100. The normalized relative expression of BmIMPM was visualized using R v4.0 with ggplot2 (https://ggplot2.tidyverse.org/).

Quantification and statistical analysis

The number of surviving H. magnanima injected with either Hm-oscar, wmks, or GFP in each replicate were counted at the adult stage. The male ratios (number of adult males/numbers of all adults) under all conditions were compared using the Steel– Dwass test in R v4.0. The sex ratio bias was also assessed based on the total numbers of (i) male and female adults and (ii) male and female embryos under each condition using the binomial test in R v4.0.

Gene expression levels between chromosomes (RNA-seq), a Z-linked Hmtpi gene expression in H. magnanima (RT-qPCR), and normalized BmIMPM expression in BmN-4 cells under all conditions (RT-qPCR) were compared using the Steel–Dwass test in R v4.0.

A phylogenetic tree of Masc proteins used in this study was constructed based on maximum likelihood with bootstrap re-sampling of 1,000 replicates using MEGA7 (Kumar et al., 2016).